```python
# Initialize Otter
import otter
grader = otter.Notebook("check_bayes_filter.ipynb")
```

# Bayes filter using bins

This Jupyter notebook file imports (and tests) all of the code needed to do the Bayes' filter assignment. Note that the actual code is in the .py files in this directory.

For each block, go to the corresponding .py file (you can use Jupyter Notebook to edit those files, or whaterver text editor you want) and fill in the missing code. Note that there's some code that you will NOT fill in (yet) - that is for the Kalman/particle filter assignments. 

Slides for this assignment: https://docs.google.com/presentation/d/1BClQb3jZRV6kE4TME6tfnjLrWVGh-BKZ0SZ8jdTVuc8/edit?usp=sharing

## Robot ground truth
This gets all the code in Robot_ground_truth.py. If you change any code in that file then you will need to re-execute this import.

```python
import sys; sys.path.insert(0, '.')

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

```python
from robot_ground_truth import RobotGroundTruth, test_discrete_move_functions
```

```python
# Syntax check of your code
robot_gt = RobotGroundTruth()
robot_gt.set_move_left_probabilities(0.3, 0.5)
ret_value = robot_gt.move_left(0.1)
if 0.49 < ret_value < 0.61:
    print(f"Robot ground truth: Passed move left syntax check")
```

```python
robot_gt.reset_location()
robot_gt.set_move_right_probabilities(0.2, 0.1)
ret_value = robot_gt.move_right(0.1)
if 0.49 < ret_value < 0.61:
    print(f"Robot ground truth: Passed move right syntax check")
```

```python
# Edit the code in robot_ground_truth.py (the methods tagged with Bayes assignment)
test_discrete_move_functions()
```

```python
grader.check("robotGT")
```

## Robot sensors

```python
from world_ground_truth import WorldGroundTruth, test_world_ground_truth
from robot_sensors import RobotSensors, test_discrete_sensors
```

```python
# Syntax check  
robot_gt = RobotGroundTruth()
world_gt = WorldGroundTruth()
robot_sensor = RobotSensors()

probs_see_door = (0.7, 0.2)
robot_sensor.set_door_sensor_probabilites(probs_see_door[0], probs_see_door[1])
ret_value = robot_sensor.query_door(robot_gt, world_gt)
if ret_value is True or ret_value is False:
    print("Passed robot sensor syntax check")
```

```python
# No code to edit for this one - just make sure the world works properly before testing the sensors
test_world_ground_truth()

# Edit the code in robot_sensors.py (the methods tagged with Bayes assignment)
test_discrete_sensors()
```

```python
grader.check("robot_sensors")
```

## Bayes filter

This is the part where you actually *do* the bayes' filter update. There's two parts, the sensor update and the move update. There are separate test methods for each, followed by a call to the method that does both.

Edit the four methods in Bayes_filter.py

```python
from bayes_filter import BayesFilter, test_bayes_filter_sensor_update, test_move_one_direction, test_move_update
```

```python
# Syntax checks    
n_doors = 2
n_bins = 10
world_ground_truth = WorldGroundTruth()
world_ground_truth.random_door_placement(n_doors, n_bins)
robot_sensor = RobotSensors()
bayes_filter = BayesFilter()
robot_ground_truth = RobotGroundTruth()

# Syntax check 1, reset probabilities
bayes_filter.reset_probabilities(n_bins)

# Syntax check 2, update sensor
bayes_filter.update_belief_sensor_reading(world_ground_truth, robot_sensor, True)

# Syntax check 3, move
bayes_filter.update_belief_move_left(robot_ground_truth)
bayes_filter.update_belief_move_right(robot_ground_truth)

# Syntax check 4, full update
bayes_filter.one_full_update(world_ground_truth, robot_ground_truth, robot_sensor, "move_left", True)

```

```python
test_bayes_filter_sensor_update()
test_move_one_direction()
test_move_update()
```

```python
grader.check("bayes_filter")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, bayes filter. Include just this .ipynb file and ALL of the .py files

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for bayes_filter (points: 15)
```
>>> assert(test_bayes_filter_sensor_update(b_print=False))
```
```
>>> assert(test_move_one_direction(b_print=False))
```
```
>>> assert(test_move_update(b_print=False))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for robotGT (points: 5)
```
>>> assert(test_discrete_move_functions(b_print=False))
```


Tests for robot_sensors (points: 10)
```
>>> assert(test_discrete_sensors(b_print=False))
```
