```python
# Initialize Otter
import otter
grader = otter.Notebook("check_kalman_filter.ipynb")
```

# Kalman filter 

This Jupyter notebook file imports (and tests) all of the code needed to do the Kalman' filter assignment. Note that the actual code is in the .py files in this directory.

This is where you'll fill in the missing code from the last assignment

Slides for this assignment: https://docs.google.com/presentation/d/1Q6w-vczvWHanGDqbuz6H1qhTOkSrX54kf1g8NgTcipQ/edit?usp=sharing

## Robot ground truth
This gets all the code in Robot_ground_truth.py. If you change any code in that file then you will need to re-execute this import.

```python
import sys; sys.path.insert(0, '.')

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

```python
from robot_ground_truth import RobotGroundTruth, test_continuous_move_functions
```

```python
# Syntax check
robot_gt = RobotGroundTruth()
robot_gt.reset_location()
robot_gt.set_move_continuos_probabilities(0.01)
ret_amt_moved = robot_gt.move_continuous(0.0)
if -0.1 < ret_amt_moved < 0.1:
    print(f"Robot ground truth: Passed move continuous syntax check")
```

```python
# Edit the code in robot_ground_truth.py (the methods tagged with Kalman assignment)
test_continuous_move_functions()
```

```python
grader.check("robotGT")
```

## Robot sensors

```python
from robot_sensors import RobotSensors, test_continuous_sensor
```

```python
# Syntax check
robot_gt = RobotGroundTruth()
robot_sensor = RobotSensors()
robot_sensor.set_distance_wall_sensor_probabilities(sigma=0.01)
dist_with_noise = robot_sensor.query_distance_to_wall(robot_gt)
if 0.0 < dist_with_noise < 1.0:
    print("Dist wall sensor, passed syntax test")
```

```python
# Edit the code in robot_sensors.py (the methods tagged with Bayes assignment)
test_continuous_sensor()
```

```python
grader.check("robot_sensors")
```

## Kalman filter

This is the part where you actually *do* the kalman filter update. There's two parts, the sensor update and the move update. The test file runs a sensor, then a move, then a sequence of sensor, sequence of move, then a mix.

```python
from kalman_filter import KalmanFilter, test_kalman_update
```

```python
# Syntax checks
kalman_filter = KalmanFilter()
robot_ground_truth = RobotGroundTruth()
robot_sensor = RobotSensors()

# Set mu/sigmas
sensor_noise = {"mu": 0.0, "sigma": 0.1}
move_error = {"mu": 0.0, "sigma": 0.05}
robot_ground_truth.set_move_continuos_probabilities(move_error["sigma"])
robot_sensor.set_distance_wall_sensor_probabilities(sensor_noise["sigma"])

kalman_filter.update_gauss_sensor_reading(robot_sensor, 0.1)
kalman_filter.update_continuous_move(robot_ground_truth, 0.0)
```

```python
test_kalman_update()
```

```python
grader.check("kalman_filter")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, kalman filter. Include just this .ipynb file and ALL of the .py files

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for kalman_filter (points: 15)
```
>>> assert(test_kalman_update(b_print=False))
```


Tests for robotGT (points: 5)
```
>>> assert(test_continuous_move_functions(b_print=False))
```


Tests for robot_sensors (points: 10)
```
>>> assert(test_continuous_sensor(b_print=False))
```
