```python
# Initialize Otter
import otter
grader = otter.Notebook("dijkstra.ipynb")
```

# Instructions

This assignment implements Dijkstra's shortest path on a graph, finding an unvisited node in a graph, picking which one to visit, and taking a path in the map and generating waypoints along that path

Given to you:
- Priority queue
- Image handling
- Eight connected neighbors

Slides: https://docs.google.com/presentation/d/1XBPw2B2Bac-LcXH5kYN4hQLLLl_AMIgoowlrmPpTinA/edit?usp=sharing


```python
# The ever-present numpy and matplot lib
import numpy as np

# Our priority queue
import heapq

# Using imageio to read in the image
import imageio

# For reading yaml files
import yaml
```

```python
import sys; 
sys.path.insert(0, '.')

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

# Run Dijkstra on a pretend SLAM map



```python
from path_planning import convert_image, dijkstra, plot_with_path, open_image
```

```python
# Opens and threshold the SLAM map image
im, im_thresh = open_image("SLAM_map.png")

robot_start_loc = (200, 150)
# Closer one to try - change this after it works...
robot_goal_loc = (315, 250)
#robot_goal_loc = (615, 850)
zoom = 0.8

"""
print(f"Image shape {im_thresh.shape}")
for i in range(0, im_thresh.shape[1]-1):
    for j in range(0, im_thresh.shape[0]-1):
        if is_free(im_thresh, (i, j)):
            print(f"Free {i} {j}")
"""
path = dijkstra(im_thresh, robot_start_loc, robot_goal_loc)
plot_with_path(im, im_thresh, zoom=zoom, robot_loc=robot_start_loc, goal_loc=robot_goal_loc, path=path)

```

```python
# Opens and threshold a real SLAM map image
im, im_thresh = open_image("map.pgm")

robot_start_loc = (1940, 1953)
robot_goal_loc = (2135, 2045)
zoom = 0.1

path = dijkstra(im_thresh, robot_start_loc, robot_goal_loc)
plot_with_path(im, im_thresh, zoom=zoom, robot_loc=robot_start_loc, goal_loc=robot_goal_loc, path=path)
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (create"s a set)
worked_with_names = {"not filled out"}
# List of URLS TAF24 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit this .ipynb file and path_planning.py through gradescope, Lab 3: A star/Dijkstra
- You do NOT need to submit the data files - we will supply those
- Take out/suppress all print statements

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely failure is because your code ran very, very slowly and the autograder timed out.

Third reason for failure is you forgot to include **path_planning.py** and/or you put the file in a folder

Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
