```python
# Initialize Otter
import otter
grader = otter.Notebook("door_open_closed.ipynb")
```

# What is in this file/notebook

Slides: https://docs.google.com/presentation/d/1vZC32UCamhyJWJBQIuP5AXKK896mBges_eO2H3Vo3q4/edit?usp=sharing

Step 1) How do you use numpy's stats package to generalize a single random variable for a Booleans variable (T/F)?

Step 2) Implement a door sensor (which needs two T/F random variables)

Step 3) Implement a robot action (which needs FOUR opened/closed random variables)

All of the actual functional code is in **door_open_closed_classes.py**. The cells here (mostly) just functions that call methods on those classes. The reason this is implemented this way is so that you can "swap out" your own version of doing things if you really wanted to. The exception to this is the first problem, which is a stand-alone implementation of a Boolean variable

Note that most of the tests here also appear in the Python file so once the Python file passes all the assert tests then this one should, too (again, excepting the first problem). 

```python
# The imports you'll need - numpy has both numpy arrays and, in the random library, the method for 
# randomly generating a numer between zero and one (np.random.uniform)
import numpy as np

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

## Boolean random variable

TODO: Practice implementing a Boolean random variable

A boolean random variable returns True or False, with some probability. 

Since probability of returning False is (1 - probability of returning True), you only need to specify the probability of returning True.

```python
def sample_boolean_variable(prob : float):
    """ Generate one sample from a boolean variable
    @param prob contains the probability of the sensor returning True
    @returns True or False """

    # Probabilities have to be between 0 and 1...
    assert 0.0 <= prob <= 1.0

    # First, use random.uniform to generate a number between 0 and one. Note that this is a uniform distribution, not
    #  a Gaussian one. Please use np.random.uniform() for this assignment. If you find a different library that also generates
    #  random numbers (eg, mentioned in ChatGPT or Stack Exchange) I do not guarantee that the tests will work.
    zero_to_one = np.random.uniform()

    # See slides - if the number generated (zero_to_one) is below the probability of returning true, return true. Otherwise, return false
    # TODO: Return True or False, depending on value of zero_to_one
    ...
```

```python
# First, check that you have no obvious syntax errors. This does not mean that your code does the correct thing,
#. just that it has no Python errors
prob_ret_true = 0.6

ret_val = sample_boolean_variable(prob_ret_true)
if ret_val is True or ret_val is False:
    print("sample_boolean: Passed syntax check")
```

```python
# Test function
def test_boolean(test_prob_value=0.6, n_samples=1000, b_print=True):
    """ Check if the sample_boolean is doing the right thing by calling it lots of times
    @param test_prob_value - any value from 0.0 to 1.0
    @param b_print - whether or not to print out intermediate results
    @returns True if sample_boolean_variable has the correct distribution of Trues/Falses"""

    if b_print:
        print("Testing boolean")

    # Just call lots of time and count
    count_true = 0
    count_false = 0
    for _ in range(0, n_samples):
        if sample_boolean_variable(test_prob_value) == True:
            count_true += 1
        else:
            count_false += 1

    # Percentage of times the function returned True
    perc_true = count_true / (count_true + count_false)
    if b_print:
        print(f"Out of {n_samples}, {count_true} were True, and {count_false} were False")
        print(f"Perc true from sampling: {perc_true}, expected {test_prob_value}")

    if not np.isclose(perc_true, test_prob_value, atol=0.05):
        if b_print:
            print("Failed")
        return False

    if b_print:
        print("Passed")
    return True
```

```python
# Note: I have set the random seed generator to use a sequence that should return true and print out passed, with
#. 581 True and 419 False
#   --but I have seen a few cases where the random number generator has changed between versions of python/numpy.
#.  -- it should still work as long as it's close to the same
#. My output (Python library 3.11.7 (main, Dec 15 2023, 12:09:04) [Clang 14.0.6 ], numpy 1.26.4) is 581 True, 419 False
import sys
np.random.seed = 5.0
print(f"Python library {sys.version}, numpy {np.__version__}")
print(f"Boolean result: {test_boolean()}")
```

```python
grader.check("boolean_sample")
```

# Door state

Create a "door state" variable. Note, this is set up to create an instance of **DoorGroundTruth** and check that it is creating the class correctly. If you really, really want to "do your own thing", you can change the data structure the functions return/accept. Do NOT change the value passed in to **door_open**. 

TODO: Edit the **__init__** and **get_door_state** methods in **DoorGroundTruth** (search for part 1)

```python
# Re-execute this cell if you change DoorGroundTruth
from door_open_closed_classes import DoorGroundTruth
```

```python
# Functions used to test DoorGroundTruth. These just create an instance of DoorGroundTruth then return the
#. door state value (method get_door_state()). These are what are called "wrappers" - they wrap functionality
# 
def make_door_ground_truth(door_open : bool):
    # Create an instance of DoorGroundTruth
    return DoorGroundTruth(door_open)

def get_door_ground_truth(my_door_ground_truth : DoorGroundTruth):
    # Assumes my_door_ground_truth is an instance of DoorGroundTruth
    return my_door_ground_truth.get_door_state()
```

```python
# Create to instances and check that the __init__ function and the get_door_state() method are correctly implemented

# Two instances, one initialized with the door open (True) and one with the door closed (False)
door_start_open = make_door_ground_truth(True)
door_start_closed = make_door_ground_truth(False)

# You can print variables out, btw - or look in the variable window to see what the values are
#  Note: This will only work correctly when you've implemented get_door_state
print(f"Door state open: {door_start_open}")
print(f"Door state closed: {door_start_closed}")

# Use the functions in the previous cell to check that you've implemented __init__ and get_door_state correctly
assert get_door_ground_truth(door_start_open) == True
assert get_door_ground_truth(door_start_closed) == False

```

```python
grader.check("door_state")
```

# Door sensor

TODO: Implement the door sensor. The code you need to write is all in the class **DoorSensor** in **door_open_closed_classes.py**. The TODOs are labeled with "part 2"

```python
# Re-execute this cell if you change DoorSensor
from door_open_closed_classes import DoorGroundTruth, DoorSensor
```

```python
# These functions just call the class methods
#. I made the decision here to NOT pass in the prob_xx values to the __init__ values, but to have an explicit set method.
#.   This is to more clearly highlight that you need to make two random variables, (one for if the door is open, one for
#.    if the door is cloesd) and those probability values can be DIFFERENT from one another.
def make_sensor(prob_true_if_open: float, prob_false_if_closed: float):
    # Make an instance of a door sensor with uniform (0.5) probability of returning True or False
    door_sensor = DoorSensor()
    # Now change the probabilities for the random variable used if the door is open
    door_sensor.set_return_true_if_open_probability(prob=prob_true_if_open)
    #. ... and the random variable for if the door is closed. 
    door_sensor.set_return_false_if_closed_probability(prob=prob_false_if_closed)
    # Return the instance
    return door_sensor

def sample_sensor(door_gt : DoorGroundTruth, door_sensor : DoorSensor):
    """ Given the door state, make one call to the sensor sample function
    @param door_gt - an instance of DoorGroundTruth
    @param door_sensor - an instance of DoorSensor
    @ return True or False """
    # Reminder that the door sensor needs the actual ground truth of whether or not the
    #.  door is open (or closed)
    return door_sensor.sample_sensor(door_gt)
```

```python
# This is JUST a syntax check - does not check if the code does the correct thing
door_sensor = make_sensor(prob_true_if_open=0.8, prob_false_if_closed=0.7)

ret_val = sample_sensor(door_gt=door_start_open, door_sensor=door_sensor)
```

```python
# Test one combo function - check if the door sensor works properly
def test_combo(prob_true_if_open: float, prob_false_if_closed: float):
    # Make the sensor
    door_sensor = make_sensor(prob_true_if_open=prob_true_if_open, prob_false_if_closed=prob_false_if_closed)

    # Do both an open and a closed door
    n_total = 1000
    for b_door_gt, exp_val in zip([True, False], [prob_true_if_open, 1.0 - prob_false_if_closed]):
        # Make the door, with door either open or closed
        door_gt = make_door_ground_truth(door_open=b_door_gt)

        count_true = 0
        for _ in range(0, n_total):
            # Just count the true values
            if sample_sensor(door_gt, door_sensor):
                count_true += 1

        if not np.isclose(count_true / n_total, exp_val, atol=0.05):
            print(f"For door in state {door_gt}, expected {exp_val}, got {count_true / n_total}")
            return False

    return True
```

```python
# Check prob true if open
assert test_combo(prob_true_if_open=0.8, prob_false_if_closed=0.5)

# Check prob false if closed
assert test_combo(prob_true_if_open=0.5, prob_false_if_closed=0.7)

# Check both
assert test_combo(prob_true_if_open=0.7, prob_false_if_closed=0.8)
```

```python
grader.check("door_sensor")
```

## Door action - opening or closing the door

TODO: Try to open (or close) the door. The assumption is that the robot tries to open (or close) the door and the door may change to the desired state - or the action may fail. This is the transition diagram.

This is multiple boolean random variables. Rather than have one set method for each possible random variable, the **set_probabilities** method takes in the starting state, the action, the ending state, and the probability that from that starting state you would take that action and end up in that state. 

The recommendation is to store all those probabilities in some sort of matrix/list of lists. The default probability value(s) should be 0.5 (uniform probability).

TODO: Fill in the **set_probability** method and the two "sample" methods (open and close actions)

```python
# putting this here again so you can easily run it
from door_open_closed_classes import DoorGroundTruth
```

```python
# Wrapper functions
def set_probability(door_gt: DoorGroundTruth, 
                    door_initial_state : bool, 
                    action : str, 
                    door_final_state : bool, 
                    prob : float):
    door_gt.set_probability(door_initial_state=door_initial_state,
                            action=action,
                            door_final_state=door_final_state,
                            prob=prob)
    
def try_action(door_gt: DoorGroundTruth, action : str):
    assert action in DoorGroundTruth.actions
    if action == "Open":
        door_gt.robot_tries_to_open_door()
    elif action == "Close":        
        door_gt.robot_tries_to_close_door()

    return door_gt.get_door_state()
```

```python
# This is (mostly) a syntax check, but by forcing the probabilities to be one, it means the action should always succeed
# TODO part 3: try changing prob = 0.0 and check that the return state is NOT the end state
for door_start_state in [True, False]:
    for action in DoorGroundTruth.actions:
        for door_end_state in [True, False]:
            # Make a door with the given start state
            door_gt = make_door_ground_truth(door_open=door_start_state)

            # Set the probabilities for this combination
            set_probability(door_gt=door_gt, 
                            door_initial_state=door_start_state,
                            action=action,
                            door_final_state=door_end_state,
                            prob=1.0)
            # Try the action - should succeed, since we set the probability to 1
            ret_state = try_action(door_gt=door_gt, action=action)

            # Did it succeed?
            assert ret_state == door_end_state

```

```python
# Basically the same as above, except sample the action multiple times and see if the distribution
#. of resulting states is correct
def test_action(n_samples=100):
    b_ret = True
    for door_start_state in [True, False]:
        for action in DoorGroundTruth.actions:
            for door_end_state in [True, False]:
                for prob in [0.0, 0.2, 0.8, 1.0]:
                    count_in_state = 0
                    for _ in range(0, n_samples):
                        # Make a new one each time, so my_door is always in the correct starting state
                        my_door = make_door_ground_truth(door_open=door_start_state)
                        set_probability(door_gt=my_door,
                                        door_initial_state=door_start_state,
                                        action=action,
                                        door_final_state=door_end_state,
                                        prob=prob)
                        # Do the action
                        try_action(door_gt=my_door, action=action)
                        # Check the end state
                        if my_door.get_door_state() == door_end_state:
                            count_in_state += 1                    

                    prob_in_state = count_in_state / n_samples
                    if not np.isclose(prob_in_state, prob, atol=0.1):
                        print(f"Failed start state {door_start_state}, door end state {door_end_state}, action {action}")
                        print(f" Expected {prob}, got {prob_in_state}")
                        b_ret = False
    return b_ret

```

```python
b_ret = test_action()
assert b_ret
```

```python
grader.check("door_action")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# TODO - set to correct value               
# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS FW25(creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit this .ipynb file and the .py file through gradescope, Door open closed assignment
- Take out/suppress all print statements

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing



Tests for boolean_sample (points: 1)
```
>>> assert test_boolean(0.3, n_samples=10000, b_print=False)
```


Tests for door_action (points: 5)
```
>>> assert test_action()
```


Tests for door_sensor (points: 4)
```
>>> assert test_combo(prob_true_if_open=0.85, prob_false_if_closed=0.72)
```


Tests for door_state (points: 1)
```
>>> assert get_door_ground_truth(door_start_open) == True
```
```
>>> assert get_door_ground_truth(door_start_closed) == False
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
